# Assignment 3

## Q1
Given a number *n*, and a function ```f1()```, you must complete the function to return a list from [1,n]

<!-- grade_id: cell-6d5c9250b66244b0, points: 0, grade: false, solution: true -->
```python
def f1(n):
    ### BEGIN SOLUTION
    return list(range(1,n+1))
    ### END SOLUTION
    pass
```

<!-- grade_id: cell-561ce6b19c552aff, points: 1, grade: true, solution: false -->
```python
#Tests
assert f1(2)==[1,2]
assert f1(3)==[1,2,3]
### BEGIN HIDDEN TESTS
assert f1(1)==[1]
assert f1(40)==list(range(1,41))
### END HIDDEN TESTS
```

## Q2
Write a function f2(), which checks whether a number n is prime or not, it must return ```True``` if prime, else ```False``` if not a prime. <br>Note: 1 is considered to be Non-Prime

<!-- grade_id: cell-4bbfc9749c7fc9d7, points: 0, grade: false, solution: true -->
```python
def f2(n):
    ### BEGIN SOLUTION
    if(n==1):
        return False
    for i in range(2, n):
        # check for factors
        if (n % i) == 0:
            return False
    return True
    ### END SOLUTION
    pass
```

<!-- grade_id: cell-826ec4211109b15c, points: 1, grade: true, solution: false -->
```python
#Tests
assert f2(2)==True
assert f2(4)==False
### BEGIN HIDDEN TESTS
assert f2(1)==False
assert f2(6911)==True
assert f2(323)==False
### END HIDDEN TESTS
```

## Q3
HackTheBox never gets tired of encoding and decoding strings, here is another question, this one is a little straight forward <br><br> *The string was splitted into a list such that every char was an element in list, then some random chars were added in the beginning and in the end such that the original character stayed in the middle* <br><br>
You are given a function *f3*() and the list *l*, return the original string

<br><code>Ex1: if the list is ['abc','pkonh','aqwetmkoi'], then the middle char of each element is 'b','o','t' repectively hence the original string is "bot"</code>

<!-- grade_id: cell-61aba2b1b33f8ef7, points: 0, grade: false, solution: true -->
```python
def f3(l):
    ### BEGIN SOLUTION
    p=""
    for i in l:
        p+=i[len(i)//2]
    return p
    ### END SOLUTION
    pass
```

<!-- grade_id: cell-6240053ada5a86d4, points: 1, grade: true, solution: false -->
```python
#Tests
assert f3(['abc','pkonh','aqwetmkoi'])=="bot"
assert f3(['abc','pkonh','o','pkp'])=="book"
### BEGIN HIDDEN TESTS
assert f3(['xxsyy',
 'xxxpyyy',
 'xxxxxxxxxayyyyyyyyy',
 'xxxxxxxcyyyyyyy',
 'xxxeyyy',
 'xoy',
 'xxxxxnyyyyy',
 'xxxxxxoyyyyyy',
 'xxxxxxxvyyyyyyy',
 'xxxxxxxxxxxxayyyyyyyyyyyy'])=="spaceonova"
assert f3(['avavavavavavavavavavavavsyxyxyxyxyxyxyxyxyxyxyxyx',
 'avavavavavavavavavavavavavavavavavavavavavavavavavavavavavavavavavoyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyx',
 'avavavavavavavavavavavavavavavavavavavuyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyxyx',
 'avavavavavavavtyxyxyxyxyxyxyx',
 'avavavavaviyxyxyxyxyx',
 'avavavavavavavavavavavavkyxyxyxyxyxyxyxyxyxyxyxyx'])=="soutik"
# assert f3(323)==False
### END HIDDEN TESTS
```

## Q4
This question prepares you for the coming question<br>
Make a list of frequencies of characters in a string ```s``` provided the string contains only lower cases alphabets.<br>
There is a built-in function in python knows an ```ord()``` which returns the ASCII value of a char.<br>
For example, the ascii value of ```a``` is 97 hence <code>ord('a')=97</code><br>Now,this would guide you

<!-- grade_id: cell-f3ccda9a2c873a97, points: 0, grade: false, solution: true -->
```python
def f4(s): #s is string
    # As there are 26 small case alphabets, create a list having 26 elements and all to 0
    l= None
    ### BEGIN SOLUTION
    l=[0]*26
    ### END SOLUTION
    
    # Now let us read every char and make sure 'a' is at index 0, 'b' at 1 and so on.....
    # Use ord()
    # Suppose you need to access 'd' whose ascii is 100 but must be stored at index 3, we would simply subtract the
    # ascii value of 'a' which is 97, now implement yourself
    
    ### BEGIN SOLUTION
    for i in s:
        l[ord(i)-ord('a')]+=1
    ### END SOLUTION
    return l
```

<!-- grade_id: cell-ec0f891982e359ef, points: 1, grade: true, solution: false -->
```python
#Tests
assert f4('abc')==[1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
assert f4('aabbcc')==[2, 2, 2, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
### BEGIN HIDDEN TESTS
assert f4('jhgfvuyzzxkdcgajduagucgsdiun')==[2, 0, 2, 3, 0, 1, 4, 1, 1, 2, 1, 0, 0, 1, 0, 0, 0, 0, 1, 0, 4, 1, 0, 1, 1, 2]
assert f4("ajhdgjybegedhghbdybadagnfbhvfbpqoiuwsjiqzzzzuoqpp")==[3, 5, 0, 4, 2, 2, 4, 4, 2, 3, 0, 0, 0, 1, 2, 3, 3, 0, 1, 0, 2, 1, 1, 0, 2, 4]
### END HIDDEN TESTS
```

## Q5 Optional
Now, let us come to the real question <br>

HackTheBox really loves strings a lot but they likes palindromic strings more. Today, they took two strings <b>A</b> and <b>B</b>, each consisting of lower case alphabets.

They are eager to know whether it is possible to choose some <b>non empty</b> strings <b>s1</b> and <b>s2</b> where <b>s1 is a substring of A</b>, <b>s2 is a substring of B</b> such that the <b>concatenation of s1 and s2 is a palindromic string.</b><br><br>
You will be given two string A and B and return ```True``` if it possible to choose such strings <b>s1 & s2</b> else ```False```
<div style="text-align: right">-- 3 points </div>

<!-- grade_id: cell-7b8c5a1796c5b1fe, points: 0, grade: false, solution: true -->
```python
def f5(A,B): #A and B are strings
    # You can use f4()
    ### BEGIN SOLUTION
    z=f4(A)
    z1=f4(B)
    for i in range(26):
        if(z[i] and z1[i]):
            return True
    return False
    ### END SOLUTION
    pass
```

<!-- grade_id: cell-33cb3c0010494a5b, points: 3, grade: true, solution: false -->
```python
#Tests
assert f5('abc','abc')==True
assert f5('a','b')==False
assert f5('abba','baab')==True
### BEGIN HIDDEN TESTS
assert f5('a','a')==True
assert f5('yjgdefywefuyf','uaredyqtwud')==True
assert f5('qwertyuiop','asdfghjklkjhgfdsasdfghjklkmnbvcx')==False
assert f5('zxcvbnmnbvcxzxcvbnmnbvcxxcvbnm','asdfghjhgfdsndfghjjhgfdsdfghj')==True
### END HIDDEN TESTS
```

## Q6
You are given two integers a and b.

In one move, you can choose some integer *k* from *1 to 10* and <b>add it to a or subtract it from a</b>. <br><br>In other words, you choose an integer k∈[1:10] and perform ```a=a + k``` or ```a = a − k```. You may use different values of k in different moves.

Your task is to find the minimum number of moves required to obtain <em><b>b</b></em> from <em><b>a</b></em>.<br>

<div style="font-size: 85%;">PS: You may not need loops, may not</div>

<!-- grade_id: cell-c7e374efc2711302, points: 0, grade: false, solution: true -->
```python
def f6(a,b):
    ### BEGIN SOLUTION
    return ((abs(a - b) + 9) // 10)
    ### END SOLUTION
    pass
```

<!-- grade_id: cell-24536de406e7c539, points: 1, grade: true, solution: false -->
```python
# Tests
assert f6(5,5)==0    #nothing required
assert f6(13,42)==3  #13→23→32→42  (add 10, add 9, add 10).
assert f6(18,4)==2   #18→10→4  (subtract 8, subtract 6)
### BEGIN HIDDEN TESTS
assert f6(1337,420)==92
assert f6(123456789,1000000000)==87654322
assert f6(100500,9000)==9150
assert f6(81773,12)==8177
### END HIDDEN TESTS
```

```python

```