# Assignment 2

Reverse the string using slicing and store in ans1

<!-- grade_id: cell-0299b3edb1134423, points: 0, grade: false, solution: true -->
```python
s ='hello'
ans1= None
### BEGIN SOLUTION
ans1=s[::-1]
### END SOLUTION
```

<!-- grade_id: cell-3331bbdd208eb488, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert ans1=="hello"[::-1]
### END HIDDEN TESTS
```

HackTheBox Foundation came up with an interesting thing, You are given a string *p* and it is coded as follows:

*An original string ans2 was taken and every char in 'ans2' is suffixed (added in end) x number of random chars and stored in a variable p* <br><br>
Find ans2
<div style="text-align: right">-- 2 points </div>

<!-- grade_id: cell-bcbb56baf111b1cd, points: 0, grade: false, solution: true -->
```python
p='clbzwaolcnmonwy pbggrx rrpbyqmaxheo twhqxzsyitzf cjlxnyoqbpso urbtupvtcs cfhdodmkyjwim fuddkbhmr soufbiqvjhetgfume'
ans2=None
### BEGIN SOLUTION
ans2=p[::6]
### END SOLUTION
```

<!-- grade_id: cell-d91e156166669e98, points: 2, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert ans2=="congrats you did it"
### END HIDDEN TESTS
```

You are given two lists x and y. Sort x in ascending order and sort y in descending order. Then store the concatenation in ans3

<!-- grade_id: cell-4ce105b074ba9ef1, points: 0, grade: false, solution: true -->
```python
x=[3,5,1,9,7,4]
y=['limit','tends','to','infinity']
ans3=None
### BEGIN SOLUTION
x.sort()
y.sort(reverse=True)
ans3=x+y
### END SOLUTION
```

<!-- grade_id: cell-db13bb8ec526dea3, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
x=[3,5,1,9,7,4]
y=['limit','tends','to','infinity']
x.sort()
y.sort(reverse=True)
r_ans=x+y
assert ans3==r_ans
### END HIDDEN TESTS
```

You are given a list p.<br>
Change 3<sup>rd</sup> element in p to <code>rock</code>.<br>
Add <code>boss</code> in the end of list.<br>
Add another element <code>don't</code> after <code>you</code><br>
And store in ans4

<!-- grade_id: cell-0c36b20d9e316bdd, points: 0, grade: false, solution: true -->
```python
ans4=None
p=['you','definitely','too','inside','of me']
### BEGIN SOLUTION
p[2]='rock'
p.append('boss')
p.insert(1,"don't")
ans4=p
### END SOLUTION
```

<!-- grade_id: cell-bc5190db0237a7ea, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
l=['you','definitely','too','inside','of me']
l[2]='rock'
l.append('boss')
l.insert(1,"don't")
assert l==ans4
### END HIDDEN TESTS
```

```python
ans3
```

Output:
```
['you', "don't", 'definitely', 'rock', 'inside', 'of me', 'boss']
```

```python

```